---
title: "Using Tidyvsere to Manipulate Data"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(remotes)
library(gradethis)

library(nycflights13)
library(dplyr)

library(ggplot2)

tutorial_options(exercise.timelimit = 60, exercise.checker = gradethis::grade_learnr)
knitr::opts_chunk$set(echo = FALSE)
```


## <span style="color: teal;">Manipulating data</span>

#### Explaining how to manipulate data frames in R using tidyverse 

In this tutorial, we’re focusing on how to use a core tidyverse package to manipulate data: <span style="color: green;">**the dplyr package.**</span>

To explore dplyr’s basic data manipulation methods, we will learn how to use the following functions:


* Subset or `filter()` a data frame’s observations by picking their values 
* Pick variables or `select()` a data frame’s columns
* Add or `mutate()` variables to a data frame
* Condense or `summarise()` variables down to a single summary 


## <span style="color: teal;">Filtering or subsetting rows</span>

### 

<span style="color: purple;">To filter rows of a data frame</span>, you can use the `filter ()` function with the `filter` argument. For example, here is a downloaded data frame we’ll use: `library(nycflights13)`. 

This data frame includes 336,776 flights that departed from New York City in 2013. The data comes from the US Bureau of Transportation Statistics.


```{r table1}
flights
#> # A tibble: 336,776 x 19
#>    year month   day dep_time sched_dep_time dep_delay arr_time sched_arr_time
#>   <int> <int> <int>    <int>          <int>     <dbl>    <int>          <int>
#> 1  2013     1     1      517            515         2      830            819
#> 2  2013     1     1      533            529         4      850            830
#> 3  2013     1     1      542            540         2      923            850
#> 4  2013     1     1      544            545        -1     1004           1022
#> 5  2013     1     1      554            600        -6      812            837
#> 6  2013     1     1      554            558        -4      740            728
#> # … with 3.368e+05 more rows, and 11 more variables: arr_delay <dbl>,
#> #   carrier <chr>, flight <int>, tailnum <chr>, origin <chr>, dest <chr>,
#> #   air_time <dbl>, distance <dbl>, hour <dbl>, minute <dbl>, time_hour <dttm>
```

### 

If you wanted to <span style="color: purple;">filter</span> this data frame to only include rows of flights that only went out on January 1st, you would do the following:

````markdown
`r ''````{r filter}
filter(flights, month == 1, day == 1)
```
````

### 

When you run that line of code, dplyr executes the filtering operation and returns a new data frame. 

```{r table2}
filter(flights, month == 1, day == 1)
```


### 

All dplyr functions never alter their inputs. If you want to save the results, you must use the assignment operator `<-`

````markdown
`r ''````{r save-filter}
jan1 <- filter(flights, month == 1, day == 1)
```
````

```{r save}
jan1 <- filter(flights, month == 1, day == 1)
```

****************************************************************************

###

If you want to <span style="color: purple;">filter the data frame</span> based on more than one condition, you can use **logical operators** discussed in a previous section. 

For example, if you want to only include rows of flights that weren’t delayed (on arrival or departure) by `more than two hours`, you would do the following:

````markdown
`r ''````{r logical-filter}
filter(flights, arr_delay <= 120, dep_delay <= 120)
```
````

```{r logical}
filter(flights, arr_delay <= 120, dep_delay <= 120)
```


****************************************************************************

### <span style="color: green;">Exercise</span>


Modify the code below to <span style="color: purple;">filter</span> the `nycflights13` data frame to only include rows of flights that went out on December 25th:

```{r grade_filter, exercise = TRUE}
filter(flights, month == 1, day == 1)
```

```{r grade_filter-solution}
filter(flights, month == 12, day == 25)
```

```{r grade_filter-check}
grade_code("Good job. Don't worry, things will soon get harder.")
```


****************************************************************************

## <span style="color: teal;">Selecting or subsetting columns</span>

### 

To <span style="color: purple;">select specific columns</span>, you can use the ` select()` function. This function is useful when the dataset contains an abundance of variables. Instead of identifying the filter argument, we can specify the `select` argument to select specific columns.

### 

For example, if you want to create a data frame that only includes the `year` column, you can do the following:

````markdown
`r ''````{r select-flights}
  select(flights, year)
```
````


```{r select-exercise}
select(flights, year)
```

###

If you would like to only include both `year`, `month`, and `day` you can do the following:

````markdown
`r ''````{r select-flights}
  select(flights, year:month)
```
````

```{r select2}
select(flights, year:day)
```


###

If you would like to include **all** columns _except_ `year`, `month`, and `day` you can use the `-` sign as shown below:

````markdown
`r ''````{r select-flights}
  select(flights, -(year:day))
```
````

```{r select-flights-day}
select(flights, -(year:day))
```

******************************************************************

### <span style="color: green;">Exercise</span>

Create a new data frame called `flights_month` that <span style="color: purple;">only includes</span> the `month` column from the `flights` data frame.

```{r grade_select, exercise = TRUE}

```

```{r grade_select-solution}
select(flights, month)
```

```{r grade_select-check}
grade_code("Good job. Don't worry, things will soon get harder.")
```


Create a new data frame called `flights_new` that <span style="color: purple;">filters the data frame</span> to only include rows where `delay` is less than or equal to 2 hours and `selects()` the following <span style="color: purple;">columns</span>, `day`, `dep_time`, `sched_dep_time`, and `dep_delay`.

```{r grade_select2, exercise = TRUE}

```

```{r grade_select2-solution}
select((flights_new <- filter(flights, arr_delay <= 120, dep_delay <= 120)), day:dep_delay)


```

```{r grade_select2-check}
grade_code("Good job. Don't worry, things will soon get harder.")
```


******************************************************************

## <span style="color: teal;">Add variables to a data frame</span>

###

We can <span style="color: purple;">add new columns</span> that are functions of existing columns with the `mutate()` function.

The <span style="color: purple;">mutate()</span> function always adds new columns to the end of the dataset. 

###

For example, consider the subset selected from the flights data as shown below:

````markdown
`r ''````{r flights_sml}
  flights_sml <- select(flights, 
  year:day, 
  ends_with("delay"), 
  distance, 
  air_time
)
```
````

```{r mutate}
flights_sml <- select(flights, 
  year:day, 
  ends_with("delay"), 
  distance, 
  air_time
)

flights_sml
```

###

We can add two columms to this table using the `mutate()` function as follows:

````markdown
`r ''````{r mutate}

mutate(flights_sml,
  gain = dep_delay - arr_delay,
  speed = distance / air_time * 60
)

```
````

```{r mutate2}
mutate(flights_sml,
  gain = dep_delay - arr_delay,
  speed = distance / air_time * 60
)


```



###

If you only want to keep the new variables, use <span style="color: purple;">transmute()</span>:


````markdown
`r ''````{r transmute}
transmute(flights,
  gain = dep_delay - arr_delay,
  hours = air_time / 60,
  gain_per_hour = gain / hours
)
```
````

```{r transmute}
transmute(flights,
  gain = dep_delay - arr_delay,
  hours = air_time / 60,
  gain_per_hour = gain / hours
)

```

**************************************************

### <span style="color: green;">Exercise</span>

Using the data `flights_small` below, add the variable `speed` that is <span style="color: purple;">equal to</span> `distance / air_time * 60`:

```{r mutate1}
flights_small<-select(flights, -(year:tailnum))
flights_small

```

```{r grade_mutate, exercise = TRUE}

```

```{r grade_mutate-solution}
mutate(flights_small,
  speed = distance / air_time * 60
)


```

```{r grade_mutate-check}
grade_code("Good job. Don't worry, things will soon get harder.")
```


*************************************************************************************


## <span style="color: teal;">Summarize columns in a data frame</span>


###


In <span style="color: purple;">tidyverse</span>, you can use functions to summarise observations in a column of a data frame using the `summarise()` function and the `group_by()` function.

This changes the data analysis from the complete dataset to individual groups we select. 

###

For example, we can get the <span style="color: purple;">average delay per date</span> running the following:

````markdown
`r ''````{r summary}
by_day <- group_by(flights, year, month, day)
summarise(by_day, delay = mean(dep_delay, na.rm = TRUE))
```
````

```{r summary}
by_day <- group_by(flights, year, month, day)
summarise(by_day, delay = mean(dep_delay, na.rm = TRUE))
```

###

Based on all the <span style="color: purple;">dplyr</span> functions we've learned above, we can summarise observations in a column by another column. 

###

For example, we can explore the relationship between the distance and average delay for each location. 

### 

Prepare the data by:

1. <span style="color: purple;">Grouping</span> flights by destination

2. <span style="color: purple;">Summarizing</span> to compute distance, average delay, and number of flights

###

<span style="color: purple;">Use the code as shown below:</span>

````markdown
`r ''````{r}
by_dest <- group_by(flights, dest)
delay <- summarise(by_dest,
  count = n(),
  dist = mean(distance, na.rm = TRUE),
  delay = mean(arr_delay, na.rm = TRUE)
)

delay <- filter(delay, count > 20, dest != "HNL")

```
````

```{r summary2}
by_dest <- group_by(flights, dest)
delay <- summarise(by_dest,
  count = n(),
  dist = mean(distance, na.rm = TRUE),
  delay = mean(arr_delay, na.rm = TRUE)
)
delay <- filter(delay, count > 20, dest != "HNL")

```


```{r ggplot}
ggplot(data = delay, mapping = aes(x = dist, y = delay)) +
  geom_point(aes(size = count), alpha = 1/3) +
  geom_smooth(se = FALSE)


```

###

Rather than giving each data frame a name, we can slow down analysis using the <span style="color: purple;">pipe</span> `%>%`.

- The pipe improves the readability of code in a way that you can read left-to-right

- We can read it as a series of commands: group, then summarize, then filter

###

When using <span style="color: purple;">aggregation functions</span>, we have an `na.rm = TRUE` argument that removes missing values before computation

For example here's a summary of the flights `dep_delay mean` grouped by `year`, `month`, and `day`:

````markdown
`r ''````{r}
flights %>% 
  group_by(year, month, day) %>% 
  summarise(mean = mean(dep_delay, na.rm = TRUE))

```
````

```{r flights-summary}
flights %>% 
  group_by(year, month, day) %>% 
  summarise(mean = mean(dep_delay, na.rm = TRUE))

```


###

Suppose `missing values` represent cancelled flights. We can solve this issue by removing all cancelled flights from the dataset.

````markdown
`r ''````{r}
not_cancelled <- flights %>% 
  filter(!is.na(dep_delay), !is.na(arr_delay))

not_cancelled %>% 
  group_by(year, month, day) %>% 
  summarise(mean = mean(dep_delay))
```
````

```{r cancel}
not_cancelled <- flights %>% 
  filter(!is.na(dep_delay), !is.na(arr_delay))

not_cancelled %>% 
  group_by(year, month, day) %>% 
  summarise(mean = mean(dep_delay))

```

We now have a subset called `not_cancelled` with <span style="color: purple;">no missing values</span> from the `flights` dataset.

###


Here's a plot of the planes (<span style="color: purple;">identified by their tail number</span>) that have the highest average delays:

```{r summarize3}
delays <- not_cancelled %>% 
  group_by(tailnum) %>% 
  summarise(
    delay = mean(arr_delay)
  )

ggplot(data = delays, mapping = aes(x = delay)) + 
  geom_freqpoly(binwidth = 10)
```

###

Whenever you do an `aggregation`, include either a **count** `n()` or a count of non-missing values.


Here's a scatterplot of <span style="color: purple;">number of flights vs. average delay</span> including a count `n=n()`.

```{r count}
delays <- not_cancelled %>% 
  group_by(tailnum) %>% 
  summarise(
    delay = mean(arr_delay, na.rm = TRUE),
    n = n()
  )

ggplot(data = delays, mapping = aes(x = n, y = delay)) + 
  geom_point(alpha = 1/10)
```

```{r count1} 
delays %>% 
  filter(n > 25) %>% 
  ggplot(mapping = aes(x = n, y = delay)) + 
    geom_point(alpha = 1/10)

```

***************************************************************************

### <span style="color: green;">Exercise</span>

Using the `flights` data, find the <span style="color: purple;">summary of the flights mean</span> `distance` grouped by `year`, `origin`, `dest`, and `arr_delay`.

```{r avg, exercise = TRUE}
```


```{r flights_summary}
flights %>% 
  group_by(year, origin, dest, arr_delay) %>% 
  summarise(mean = mean(distance, na.rm = TRUE))

```
